## MISA (2024-2025)
- Alohan'ny mamerina dia avereno atao Run ny notebook iray manontolo. Ny fanaovana azy dia redémarrena mihitsy ny kernel aloha (jereo menubar, safidio **Kernel$\rightarrow$Restart Kernel and Run All Cells**).

- Izay misy hoe `YOUR CODE HERE` na `YOUR ANSWER HERE` ihany no fenoina. Afaka manampy cells vaovao raha ilaina. Aza adino ny mameno references eo ambany raha ilaina.

## References
chatGPT, Claude AI

---

```python
from random import randrange
import numpy as np
from sklearn.metrics import mean_squared_error, log_loss
from sklearn.datasets import load_breast_cancer, load_diabetes


def grad_check_sparse(f, x, analytic_grad, num_checks=10, h=1e-5, error=1e-9):
    """
    sample a few random elements and only return numerical
    in this dimensionslinear_models_cvxpy.ipynb
    """

    for i in range(num_checks):
        ix = tuple([randrange(m) for m in x.shape])

        oldval = x[ix]
        x[ix] = oldval + h  # increment by h
        fxph = f(x)  # evaluate f(x + h)
        x[ix] = oldval - h  # increment by h
        fxmh = f(x)  # evaluate f(x - h)
        x[ix] = oldval  # reset

        grad_numerical = (fxph - fxmh) / (2 * h)
        grad_analytic = analytic_grad[ix]
        rel_error = abs(grad_numerical - grad_analytic) / (
            abs(grad_numerical) + abs(grad_analytic)
        )
        print(
            "numerical: %f analytic: %f, relative error: %e"
            % (grad_numerical, grad_analytic, rel_error)
        )
        assert rel_error < error

def rel_error(x, y):
    """ returns relative error """
    return np.max(np.abs(x - y) / (np.maximum(1e-8, np.abs(x) + np.abs(y))))
```

# Linear regression

```python
data = load_diabetes()
X_train1, y_train1 = data.data, data.target
w1 = np.random.randn(X_train1.shape[1]) * 0.0001
b1 = np.random.randn(1) * 0.0001
```

<!-- grade_id: cell-f798c8cb8fdf24eb, points: 0, grade: false, solution: true -->
```python
def mse_loss_naive(w, b, X, y, alpha=0):
    """
    MSE loss function WITH FOR LOOPs
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0

    N = X.shape[0]

    for i in range(N):
        y_pred = np.dot(X[i], w) + b
        error = y_pred - y[i]

        loss += error ** 2
        dw += 2 * error * X[i]
        db += 2 * error

    loss /= N
    dw /= N
    db /= N

    loss += alpha * np.sum(w ** 2)
    dw += 2 * alpha * w

    return loss, dw, np.array(db).reshape(1,)

```

### Naive Linear regression loss

<!-- grade_id: cell-84268496cd13d9b2, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  2.502519422824196e-16
Gradient check w
numerical: -1.553189 analytic: -1.553189, relative error: 4.435811e-08
numerical: -2.801912 analytic: -2.801914, relative error: 3.267817e-07
numerical: -4.145420 analytic: -4.145419, relative error: 1.227224e-07
numerical: -4.145420 analytic: -4.145419, relative error: 1.227224e-07
numerical: 2.892060 analytic: 2.892060, relative error: 1.481810e-08
numerical: -4.145420 analytic: -4.145419, relative error: 1.227224e-07
/tmp/ipykernel_41658/1179089906.py:29: DeprecationWarning: Conversion of an array with ndim > 0 to a scalar is deprecated, and will error in future. Ensure you extract a single element from your array before performing this operation. (Deprecated NumPy 1.25.)
  "numerical: %f analytic: %f, relative error: %e"
numerical: -1.553189 analytic: -1.553189, relative error: 4.435811e-08
numerical: -1.376393 analytic: -1.376395, relative error: 5.321007e-07
numerical: -1.376393 analytic: -1.376395, relative error: 5.321007e-07
```

### Naive Ridge regression loss

<!-- grade_id: cell-a8e942346b61f083, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -0.315301 analytic: -0.315302, relative error: 1.851010e-06
numerical: -1.553843 analytic: -1.553843, relative error: 6.220918e-08
numerical: -3.153479 analytic: -3.153477, relative error: 2.933103e-07
numerical: -1.553843 analytic: -1.553843, relative error: 6.220918e-08
numerical: -3.153479 analytic: -3.153477, relative error: 2.933103e-07
numerical: 2.892066 analytic: 2.892066, relative error: 2.536572e-09
/tmp/ipykernel_41658/1179089906.py:29: DeprecationWarning: Conversion of an array with ndim > 0 to a scalar is deprecated, and will error in future. Ensure you extract a single element from your array before performing this operation. (Deprecated NumPy 1.25.)
  "numerical: %f analytic: %f, relative error: %e"
numerical: -4.295886 analytic: -4.295885, relative error: 1.309717e-07
numerical: -1.553843 analytic: -1.553843, relative error: 6.220918e-08
numerical: -1.274773 analytic: -1.274774, relative error: 2.861052e-07
numerical: -0.315301 analytic: -0.3
```

<!-- grade_id: cell-5470d6ef89b09593, points: 0, grade: false, solution: true -->
```python
def mse_loss_vectorized(w, b, X, y, alpha=0):
    """
    MSE loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    N = X.shape[0]
    
    y_pred = X @ w + b
    error = y_pred - y
    loss = np.mean(error ** 2)
    
    dw = (2 / N) * (X.T @ error)
    db = (2 / N) * np.sum(error)
    
    if alpha > 0:  
        loss += (alpha / (2 * N)) * np.sum(w ** 2) 
        dw += (alpha / N) * w 
    
    return loss, dw, np.array(db).reshape(1,)
```

### Vectorised Linear regression loss

<!-- grade_id: cell-6b4e54d894e2ebee, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  0.0
Gradient check w
numerical: 2.892060 analytic: 2.892060, relative error: 1.481810e-08
numerical: -3.153318 analytic: -3.153318, relative error: 5.088761e-08
numerical: -4.296088 analytic: -4.296088, relative error: 1.042735e-08
numerical: -0.315454 analytic: -0.315454, relative error: 2.295677e-07
numerical: -4.296088 analytic: -4.296088, relative error: 1.042735e-08
numerical: -3.234111 analytic: -3.234111, relative error: 8.533431e-09
numerical: -1.553189 analytic: -1.553189, relative error: 7.275509e-08
numerical: -1.376394 analytic: -1.376395, relative error: 1.356322e-07
numerical: -3.234111 analytic: -3.234111, relative error: 8.533431e-09
numerical: -1.553189 analytic: -1.553189, relative error: 7.275509e-08
numerical: -3.234111 analytic: -3.234111, relative error: 8.533431e-09
numerical: -1.376394 analytic: -1.376395, relative error: 1.356322e-07
numerical: -4.145419 analytic: -4.145419, relative error: 8.916107e-09
numerical: -1.553189 analytic: -1.553189, re
```

### Vectorized ridge regression loss

<!-- grade_id: cell-76aafc86c4b9d590, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -1.553190 analytic: -1.553190, relative error: 7.657678e-08
numerical: -1.376394 analytic: -1.376395, relative error: 1.187555e-07
numerical: -3.153318 analytic: -3.153318, relative error: 5.062874e-08
numerical: -1.275044 analytic: -1.275044, relative error: 8.992565e-09
numerical: -1.376394 analytic: -1.376395, relative error: 1.187555e-07
numerical: -2.801914 analytic: -2.801914, relative error: 8.279809e-09
numerical: -1.275044 analytic: -1.275044, relative error: 8.992565e-09
numerical: -1.376394 analytic: -1.376395, relative error: 1.187555e-07
numerical: 2.892060 analytic: 2.892060, relative error: 1.363025e-08
numerical: 2.892060 analytic: 2.892060, relative error: 1.363025e-08
numerical: -2.801914 analytic: -2.801914, relative error: 8.279809e-09
numerical: -4.296087 analytic: -4.296087, relative error: 5.242173e-09
numerical: -3.153318 analytic: -3.153318, relative error: 5.062874e-08
numerical: 2.892060 analytic: 2.892060, relative error: 1.363025
```

# Logistic regression

```python
def sigmoid(z):
    return 1 / (1 + np.exp(-z))

data = load_breast_cancer()
X_train2, y_train2 = data.data, data.target
w2 = np.random.randn(X_train2.shape[1]) * 0.0001
b2 = np.random.randn(1) * 0.0001
```

### Naive

<!-- grade_id: cell-879a7232cd873202, points: 0, grade: false, solution: true -->
```python
def log_loss_naive(w, b, X, y, alpha=0):
    """
    Log loss function WITH FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    N = X.shape[0]
    
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    eps = 1e-15
    
    for i in range(N):
        z = np.dot(X[i], w.flatten()) + b  
        y_pred = sigmoid(z)
        
        # avoid log(0)
        y_pred = np.clip(y_pred, eps, 1 - eps)
        
        # Binary cross-entropy loss
        loss += -(y[i] * np.log(y_pred) + (1 - y[i]) * np.log(1 - y_pred))
        
        error = y_pred - y[i]
        dw += (error * X[i])
        db += error
    
    loss /= N
    dw /= N
    db /= N
    
    if alpha > 0:
        loss += (alpha / (2 * N)) * np.sum(w ** 2)
        dw += (alpha / N) * w
    
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-edbf23327a9bc955, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  3.15420236086316e-16
Gradient check w
numerical: 0.282557 analytic: 0.282557, relative error: 1.212506e-10
numerical: 0.018888 analytic: 0.018888, relative error: 1.844187e-10
numerical: -0.016164 analytic: -0.016164, relative error: 6.507748e-10
/tmp/ipykernel_41658/1179089906.py:29: DeprecationWarning: Conversion of an array with ndim > 0 to a scalar is deprecated, and will error in future. Ensure you extract a single element from your array before performing this operation. (Deprecated NumPy 1.25.)
  "numerical: %f analytic: %f, relative error: %e"
numerical: 2.594055 analytic: 2.594055, relative error: 1.178303e-07
numerical: 61.128437 analytic: 61.128720, relative error: 2.309855e-06
numerical: 123.539027 analytic: 123.539994, relative error: 3.915864e-06
numerical: 0.282557 analytic: 0.282557, relative error: 1.212506e-10
numerical: 0.282557 analytic: 0.282557, relative error: 1.212506e-10
numerical: -0.013530 analytic: -0.013530, relative error: 4.835926e-10
numeri
```

### Naive with regulariztion

<!-- grade_id: cell-2482505d2f813068, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: 0.041274 analytic: 0.041274, relative error: 7.636450e-10
/tmp/ipykernel_41658/1179089906.py:29: DeprecationWarning: Conversion of an array with ndim > 0 to a scalar is deprecated, and will error in future. Ensure you extract a single element from your array before performing this operation. (Deprecated NumPy 1.25.)
  "numerical: %f analytic: %f, relative error: %e"
numerical: -0.104143 analytic: -0.104143, relative error: 5.623462e-09
numerical: -0.030082 analytic: -0.030081, relative error: 5.534666e-06
numerical: -1.158954 analytic: -1.158954, relative error: 2.253811e-09
numerical: -0.006306 analytic: -0.006306, relative error: 1.487702e-09
numerical: -0.013530 analytic: -0.013530, relative error: 4.157252e-10
numerical: -0.123659 analytic: -0.123659, relative error: 4.292576e-11
numerical: 0.000631 analytic: 0.000631, relative error: 1.399723e-08
numerical: -0.005492 analytic: -0.005492, relative error: 2.734794e-09
numerical: 0.018888 analytic: 0.01888
```

### Vectorized

<!-- grade_id: cell-e2043c379ba3cc3e, points: 0, grade: false, solution: true -->
```python
def log_loss_vectorized(w, b, X, y, alpha=0):
    """
    log loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    N = X.shape[0]
    
    eps = 1e-15
    
    z = X @ w + b
    y_pred = sigmoid(z)
    y_pred = np.clip(y_pred, eps, 1 - eps)
    
    loss = np.mean(-(y * np.log(y_pred) + (1 - y) * np.log(1 - y_pred)))
    
    # Vectorized gradients
    error = y_pred - y
    dw = (1 / N) * (X.T @ error)
    db = (1 / N) * np.sum(error)
    
    # regularization term (L2)
    if alpha > 0:
        loss += (alpha / (2 * N)) * np.sum(w ** 2)
        dw += (alpha / N) * w
    
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-19b05ba596d391cd, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  0.0
Gradient check w
numerical: -1.027986 analytic: -1.027986, relative error: 1.058989e-09
numerical: 0.039093 analytic: 0.039093, relative error: 1.607596e-12
numerical: 0.041274 analytic: 0.041274, relative error: 3.164845e-12
numerical: 0.041274 analytic: 0.041274, relative error: 3.164845e-12
numerical: 0.282557 analytic: 0.282557, relative error: 8.195860e-11
numerical: 123.539027 analytic: 123.539994, relative error: 3.915864e-06
numerical: -0.000805 analytic: -0.000805, relative error: 4.907863e-09
numerical: -0.007084 analytic: -0.007084, relative error: 2.654528e-10
numerical: 0.018888 analytic: 0.018888, relative error: 1.094744e-10
numerical: 0.005298 analytic: 0.005298, relative error: 3.936003e-10
numerical: -0.006306 analytic: -0.006306, relative error: 4.092361e-10
numerical: 0.041274 analytic: 0.041274, relative error: 3.164845e-12
numerical: -0.000805 analytic: -0.000805, relative error: 4.907863e-09
numerical: 0.014615 analytic: 0.014615, relative error
```

### Vectorized with regularization

<!-- grade_id: cell-0c05b3e99335f4cd, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -0.030082 analytic: -0.030081, relative error: 5.534758e-06
numerical: 0.000631 analytic: 0.000631, relative error: 7.949336e-10
numerical: 0.005298 analytic: 0.005298, relative error: 3.983735e-10
numerical: -1.158954 analytic: -1.158954, relative error: 2.265786e-09
numerical: 0.282557 analytic: 0.282557, relative error: 8.300001e-11
numerical: -0.007084 analytic: -0.007084, relative error: 2.201861e-10
numerical: 0.005298 analytic: 0.005298, relative error: 3.983735e-10
numerical: 0.000081 analytic: 0.000081, relative error: 3.965727e-08
numerical: -0.000278 analytic: -0.000278, relative error: 9.972568e-09
numerical: 8.637192 analytic: 8.637193, relative error: 3.400249e-08
numerical: 0.282557 analytic: 0.282557, relative error: 8.300001e-11
numerical: 0.282557 analytic: 0.282557, relative error: 8.300001e-11
numerical: 0.274948 analytic: 0.274948, relative error: 3.704624e-09
numerical: 0.000081 analytic: 0.000081, relative error: 3.965727e-08
numerical
```

# Gradient descent for Linear models

<!-- grade_id: cell-076e6ebd0bdf09ff, points: 0, grade: false, solution: true -->
```python
class LinearModel():
    def __init__(self):
        self.w = None
        self.b = None

    def train(self, X, y, learning_rate=1e-3, alpha=0, num_iters=100, batch_size=200, verbose=False):
        N, d = X.shape
        
        if self.w is None: # Initialization
            self.w = 0.001 * np.random.randn(d)
            self.b = 0.0

        # Run stochastic gradient descent to optimize w
        
        loss_history = []
        for it in range(num_iters):
            # Sample batch_size elements in X_batch and y_batch
            # X_batch shape is  (batch_size, d) and y_batch shape is (batch_size,)
            # Hint: Use np.random.choice to generate indices
            indices = np.random.choice(N, batch_size, replace=True)
            X_batch = X[indices]
            y_batch = y[indices]
            
            # evaluate loss and gradient
            loss, dw, db = self.loss(X_batch, y_batch, alpha)
            loss_history.append(loss)

            # perform parameter update
            # Update the weights w and bias b using the gradient and the learning rate.
            self.w = self.w - learning_rate * dw
            self.b = self.b - learning_rate * db
            
            if verbose and it % 10000 == 0:
                print("iteration %d / %d: loss %f" % (it, num_iters, loss))
                
        return loss_history

    def predict(self, X):
        pass

    def loss(self, X_batch, y_batch, reg):
        pass

class LinearRegressor(LinearModel):
    """ Linear regression """

    def loss(self, X_batch, y_batch, alpha):
        return mse_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        return X @ self.w + self.b

class LogisticRegressor(LinearModel):
    """ Logistic regression """

    def loss(self, X_batch, y_batch, alpha):
        return log_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        """ Return prediction labels vector of 0 or 1 """
        y_pred = sigmoid(X @ self.w + self.b)
        return (y_pred >= 0.5).astype(int)
```

## Linear regression with gradient descent

<!-- grade_id: cell-92f36a3b387a4277, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LinearRegression

sk_model = LinearRegression(fit_intercept=True)
sk_model.fit(X_train1, y_train1)
sk_pred = sk_model.predict(X_train1)
sk_mse = mean_squared_error(sk_pred, y_train1)

model = LinearRegressor()
model.train(X_train1, y_train1, num_iters=75000, batch_size=64, learning_rate=1e-2, verbose=True)
pred = model.predict(X_train1)
mse = mean_squared_error(pred, y_train1)

print("MSE scikit-learn:", sk_mse)
print("MSE gradient descent model :", mse)
assert mse - sk_mse < 100
```

Output:
```
iteration 0 / 75000: loss 32109.162661
iteration 10000 / 75000: loss 4002.370697
iteration 20000 / 75000: loss 3271.168236
iteration 30000 / 75000: loss 2815.141751
iteration 40000 / 75000: loss 2555.015600
iteration 50000 / 75000: loss 3253.578513
iteration 60000 / 75000: loss 2747.760111
iteration 70000 / 75000: loss 3096.703462
MSE scikit-learn: 2859.6963475867506
MSE gradient descent model : 2884.6782243151547
```

## Logistc regression with gradient descent

<!-- grade_id: cell-20c0450a0bbb2b6c, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LogisticRegression
from sklearn.preprocessing import StandardScaler

scaler = StandardScaler()
X_train2 = scaler.fit_transform(X_train2)

sk_model = LogisticRegression(fit_intercept=True)
sk_model.fit(X_train2, y_train2)
sk_pred = sk_model.predict(X_train2)
sk_log_loss = log_loss(sk_pred, y_train2)

model = LogisticRegressor()
model.train(X_train2, y_train2, num_iters=75000, batch_size=64, learning_rate=1e-3, verbose=True)
pred = model.predict(X_train2)
model_log_loss = log_loss(pred, y_train2)

print("Log-loss scikit-learn:", sk_log_loss)
print("Log-loss gradiet descent model :", model_log_loss)
print("Error :", rel_error(sk_log_loss, model_log_loss))
assert rel_error(sk_log_loss, model_log_loss) < 1e-7
```

Output:
```
iteration 0 / 75000: loss 0.692250
iteration 10000 / 75000: loss 0.143274
iteration 20000 / 75000: loss 0.047884
iteration 30000 / 75000: loss 0.101233
iteration 40000 / 75000: loss 0.025982
iteration 50000 / 75000: loss 0.088900
iteration 60000 / 75000: loss 0.121153
iteration 70000 / 75000: loss 0.051336
Log-loss scikit-learn: 0.44341928598210933
Log-loss gradiet descent model : 0.44341928598210933
Error : 0.0
```

```python

```